Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nesma Silman"
#COLLABORATORS = ""
```

---

# Information Visualization I 
## School of Information, University of Michigan

## Week 1: 
- Domain identification vs Abstract Task extraction
- Pandas Review 

## Assignment Overview
### The objectives for this week are for you to:

- Review, reflect, and apply the concepts of Domain Tasks and Abstract Tasks. Specifically, given a real context, identify the expert's goals and then abstract the visualization tasks. 

!["Drawing"](assets/domain-abstraction.png)



- Review and evaluate the domain of [Pandas](https://pandas.pydata.org/) as a tool for reading, manipulating, and analyzing datasets in Python.

### The total score of this assignment will be 100 points consisting of:
- Case study reflection: Car congestion and crash rates (20 points)
- Pandas programming exercise (80 points)


### Resources:
- We're going to be recreating parts of this article by [CMAP](https://www.cmap.illinois.gov/) available [online](https://www.cmap.illinois.gov/updates/all/-/asset_publisher/UIMfSLnFfMB6/content/crash-scans-show-relationship-between-congestion-and-crash-rates) (CMAP, 2016)  
- We'll need the datasets from the city of Chicago. We have downloaded a subset to the local folder [/assets](assets/)
    - If you're curious, the original dataset can be found on [Chicago Data Portal](https://data.cityofchicago.org/)
        - [Chicago Traffic Tracker - Historical Congestion Estimates by Segment - 2011-2018](https://data.cityofchicago.org/Transportation/Chicago-Traffic-Tracker-Historical-Congestion-Esti/77hq-huss)
        - [Traffic Crashes - Crashes](https://data.cityofchicago.org/Transportation/Traffic-Crashes-Crashes/85ca-t3if)
- Altair
    - We will use a python library called [Altair](https://altair-viz.github.io/) for the visualizations. Don't worry about understanding this code. You will only need to prepare the data for the visualization in Pandas. If you do it correctly, our code will produce the visualization for you.

## Part 1. Domain identification vs Abstract Task extraction (20 points)
Read the following article by CMAP [Crash scans show the relationship between congestion and crash rates](https://www.cmap.illinois.gov/updates/all/-/asset_publisher/UIMfSLnFfMB6/content/crash-scans-show-relationship-between-congestion-and-crash-rates) and answer the following questions:


### 1.1 Briefly describe who you think performed this analysis. What is their expertise? What is their goal for the article? Give 3 examples of domain tasks featured in the article. (10 points)

<!-- grade_id: cell-d5cc91ff23d1fd3d, points: 0, grade: true, solution: true -->
>This analysis was performed by CMAP (The Chicago Metropolitan Agency for Planning). It is Illinois region’s comprehensive planning organization.

>Their expertise is implementing strategies that address transportation and other quality-of-life issues. They developed and are now implementing ON TO 2050, a new long-range plan to achieve these strategies.

>Their goal for the article in summary is to determine the underlying problems that lead to high crash rates - with a focus on expressways - and start solving these problems to decrease the crashes rate.

>More detailed goals:
Their plan "GO TO 2040" identifies public safety as a crucial goal. This includes developing a transportation system that reduces the risk of crashes. GO TO 2040 supports strategies by transportation agencies to reduce crashes through project design, enforcement, education, and other approaches. CMAP also works with partner agencies through CMAP's Advanced Technology Task Force and Regional Transportation Operations Coalition to identify, finance, and implement needed technology and highway operational improvements to reduce the number of crashes and respond to them quickly when they do occur.


>Domain task # 1: Measuring the speed at five-minute intervals throughout a typical weekday on expressways in the Chicago region.

>Domain task # 2: The relationship between congestion and crashes on all expressway segments that had the same travel time index value, which is a unit for measuring congestion (It measures the amount of time a trip takes in the peak period relative to uncongested conditions).

>Domain task # 3: Find the number of crashes that occurred per 100 million vehicle miles traveled (VMT) between 2008 and 2012 at points one-tenth of a mile apart along almost all of the expressways and tollways in the region. (Each task handles the number of crashes per each road - one at a time).

### 1.2  For each domain task describe the abstract task (10 points)

<!-- grade_id: cell-06458778598db3cb, points: 0, grade: true, solution: true -->
>Abstract task # 1: Finding out the change of a quantitative variable over time.

>Abstract task # 2: Find the correlation between 2 Quantitative variables.

>Abstract task # 3: Find the correlation between 1 quantitative variable and 1 geographic variable, this task is repeated many times to handle this correlation over many categorical variables - one at a time.

## Part 2. Pandas programming exercise (80 points)
We have provided some code to create visualizations based on these two datasets:
1. [Historic Congestion](assets/Pulaski.small.csv.gz) 
2. [Traffic Crashes](assets/Traffic.Crashes.csv.gz)

Complete each assignment function and run each cell to generate the final visualizations


```python
import pandas as pd
import numpy as np
import altair as alt
```

```python
# enable correct rendering
alt.renderers.enable('default')
```

Output:
```
RendererRegistry.enable('default')
```

```python
# uses intermediate json files to speed things up
alt.data_transformers.enable('json')
```

Output:
```
DataTransformerRegistry.enable('json')
```

### PART A: Historic Congestion ( 55 points)
For parts 2.1 to 2.5 we will use the Historic Congestion dataset. This dataset contains measures of speed for different segments. For this subsample, the available measures are limited to traffic on Pulaski Road in 2018.

### 2.1 Read and resample (15 points)
Complete the `read_csv` and `get_group_first_row` functions.
Since our dataset is large we want to only grab one measurement per hour for each segment. To do this, we will resample by selecting the first measure for each month, day, hour on each segment. Complete the `get_group_first_row` function to achieve this. Note that the file we are loading is compressed--depending on how you load the file, this may or may not make a difference ([you'll want to look at the API documents](https://pandas.pydata.org/pandas-docs/stable/reference/index.html)).

<!-- grade_id: read_file, points: 0, grade: false, solution: true -->
```python
def read_csv(filename):
    """Read the csv file from filename (uncompress 'gz' if needed)
    return the dataframe resulting from reading the columns
    """
    fl=pd.read_csv(filename,compression='gzip')
    df=pd.DataFrame(fl)
    return df
    raise NotImplementedError()
```

<!-- grade_id: read_file_test, points: 5, grade: true, solution: false -->
```python
# Save the congestion dataframe on hist_con
hist_con = read_csv('assets/Pulaski.small.csv.gz')
print(hist_con.shape)
assert hist_con.shape == (3195450, 10)
assert list(hist_con.columns) == ['TIME','SEGMENT_ID','SPEED','STREET','DIRECTION','FROM_STREET','TO_STREET',
                                  'HOUR','DAY_OF_WEEK','MONTH']
```

Output:
```
(3195450, 10)
```

<!-- grade_id: resample, points: 0, grade: false, solution: true -->
```python
def get_group_first_row(df, grouping_columns):
    """Group rows for the grouping columns and return the first row belonging to each group
    (you can look at first() for reference). We'll write this function to be more general in case
    we want to use it for a different resample.
    return a dataframe without a hierarchical index (use default index)
    
    See the example link below if you want a better sense of what this should return
    """
    grp_df=df.groupby(grouping_columns).first()
    grp_df.reset_index(inplace=True)
    return grp_df
    raise NotImplementedError()
```

```python
# test your code, we want segment_rows to be resampled version of hist_con where we've grouped by the
# properties month, day_of_week, hour, and segment_id and returned the first measure of each group
segment_rows = get_group_first_row(hist_con, ['MONTH','DAY_OF_WEEK', 'HOUR', 'SEGMENT_ID'])
```

The table should look something like [this](assets/segment_rows.png)

<!-- grade_id: resample_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell

```

### 2.2 Basic Bar Chart Visualization (10 points)
We want to create a visualization for the *average speed* of each segment (across all the samples). To do this, we're going to want to group by each segment and calculate the average speed on each. Complete this code on the `average_speed_per_segment` function.

<!-- grade_id: average_speed_per_segment, points: 0, grade: false, solution: true -->
```python
def average_speed_per_segment(df):
    """Group rows by SEGMENT_ID and calculate the mean of each
    return a *series* where the index is the segment id and each value is the average speed per segment
    """
    spd=df.groupby('SEGMENT_ID').agg(np.mean)['SPEED']
    return spd
    raise NotImplementedError()
```

```python
# calculate the average speed per segment
average_speed = average_speed_per_segment(segment_rows)

# create labels for the visualization
labels = average_speed.index.astype(str)

# grab the values from the table
values = pd.DataFrame(average_speed).reset_index()

# check what's in average_speed
average_speed

```

Output:
```
SEGMENT_ID
19    12.251926
20    15.274452
21    12.141079
22    12.346769
23    12.716657
        ...    
93    13.503260
94    14.560759
95    14.959099
96    21.659751
97    18.714286
Name: SPEED, Length: 78, dtype: float64
```

If you got things right, the table should look something like [this](assets/average_speed.png)

<!-- grade_id: average_speed_per_segment_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell

```

```python
# let's generate the visualization

# create a chart
base = alt.Chart(values)

# we're going to "encode" the variables, more on this next assignment
encoding = base.encode(
    x= alt.X(
            'SEGMENT_ID:Q',
            title='Segment ID',
            scale=alt.Scale(zero=False)
    ),
    y=alt.Y(
            'sum(SPEED):Q',
            title='Speed Average MPH'
    ),
)

# we're going to use a bar chart and set various parameters (like bar size and title) to make it readable
encoding.mark_bar(size=7).properties(title='Average Speed per Segment',height=300, width=900)
```

Output:
```
alt.Chart(...)
```

### 2.3 Create a basic pivot table (10 points)
For the next visualization, we need a more complex transformation that will allow us to see the average speed for each month. To do this, we will create a pivot table where the index is the month, and each column is a segment id. We will put the average speed in the cells. From the table, we'll be able to find the month (by index)--giving us the row, and pick the column corresponding to the segment we care about.

Complete the `create_pivot_table` function for this

```python
toto=get_group_first_row(read_csv('assets/Pulaski.small.csv.gz'), ['MONTH','DAY_OF_WEEK', 'HOUR', 'SEGMENT_ID'])
toto.head()
```

Output:
```
MONTH  DAY_OF_WEEK  HOUR  SEGMENT_ID                    TIME  SPEED  \
0      2            4    17          19  02/28/2018 05:40:00 PM     -1   
1      2            4    17          20  02/28/2018 05:40:00 PM     15   
2      2            4    17          21  02/28/2018 05:40:00 PM     29   
3      2            4    17          22  02/28/2018 05:40:00 PM     28   
4      2            4    17          23  02/28/2018 05:40:00 PM     23   

    STREET DIRECTION FROM_STREET TO_STREET  
0  Pulaski        NB       111th     107th  
1  Pulaski        NB       107th     103rd  
2  Pulaski        NB       103rd      99th  
3  Pulaski        NB        99th      95th  
4  Pulaski        NB        95th      91st
```

<!-- grade_id: create_pivot_table, points: 0, grade: false, solution: true -->
```python
def create_pivot_table (df):
    """return a pivot table where:
    each row i is a month
    each column j is a segment id
    each cell value is the average speed for the month i in the segment j
    """
    return df.pivot_table(values='SPEED',index='MONTH',columns='SEGMENT_ID', aggfunc='mean')
    raise NotImplementedError()
```

```python
# run the code and see what's in the table
pivot_table = create_pivot_table(segment_rows)
pivot_table
```

Output:
```
SEGMENT_ID         19         20         21         22         23         24  \
MONTH                                                                          
2            6.857143  16.142857  13.571429  19.571429  18.285714  15.857143   
3           10.773810  14.863095  11.696429  11.815476  13.583333  16.244048   
4           11.744048  14.958333  11.791667  12.071429  13.208333  16.779762   
5           11.357143  14.738095  11.369048  11.916667  12.023810  13.220238   
6           11.630952  14.583333  13.011905  12.279762  12.428571  14.678571   
7           11.755952  13.595238  10.880952  12.238095  12.267857  14.321429   
8           12.988095  15.446429  12.303571  13.315476  13.023810  15.827381   
9           13.970238  17.059524  14.398810  13.452381  12.017857  14.869048   
10          13.708333  15.666667  12.434524  13.041667  12.422619  13.714286   
11          12.970238  16.107143  11.922619  11.476190  12.125000  15.607143   
12          11.845238  15.690476  11.541
```

The table should look something like [this](assets/pivot_table.png)

```python
# we're going to implement a transformation to put the pivot table into a 'long form' because it
# is easier to specify the visualization. You can print out hm_pivot_table to see what it looks like
hm_pivot_table = pivot_table.copy().unstack().reset_index()
hm_pivot_table['SPEED'] = hm_pivot_table[0]
hm_pivot_table.drop(0,axis=1,inplace=True)

# create the visualization. We're going to use rectangles (a heat map of sorts). We'll use the segment_id to
# figure out the horizontal placement (x), the month as the vertical (y) and use color to encode the speed.
encoding = alt.Chart(hm_pivot_table).mark_rect().encode(
    x='SEGMENT_ID:O',
    y='MONTH:O',
    color='SPEED:Q'
)

encoding.properties(title='Average Speed per Segment per Month',height=300, width=800)
```

Output:
```
alt.Chart(...)
```

<!-- grade_id: create_pivot_table_test, points: 10, grade: true, solution: false -->
```python
# test function
pivot_table = create_pivot_table(segment_rows)
# check that the rows are months and columns are segments
assert pivot_table.shape == (11, 78), "Problem 2.3, first test"
# check that the value is the average
assert int(pivot_table.loc[2,19]) == 6, "Problem 2.3, second test"
assert int(pivot_table.loc[3,19]) == 10, "Problem 2.3, third test"
```

### 2.4 Sorting, Transforming, and Filtering (20 points)
Without telling you too much about the visualization we want to create next (that's part of the bonus below), we need to get the data into a form we can use. 
- We're going to need to sort the dataframe by one or more columns (this is the `sort` function). 
- We'll want to create a derivative column that is the time of the measurement rounded to the nearest hour (`time_to_hours`)
- We need to "facet" the data into groups to generate different visualizations. 
- We need a function that selects part of the dataframe that matches a specific characteristic (`filter_orientation`)
- Grab a specific column from the dataframe (`select_column`)

<!-- grade_id: sort, points: 0, grade: false, solution: true -->
```python
def sort(df, sorting_columns):
    """Sort the rows by the columns
    return the sorted dataframe
    """
    df.sort_values(by=sorting_columns, inplace=True)
    return df
    raise NotImplementedError()
```

```python
segment_rows = sort(segment_rows, ['SEGMENT_ID'])
```

<!-- grade_id: sort_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell

```

<!-- grade_id: time_to_hours, points: 0, grade: false, solution: true -->
```python
def time_to_hours(df):
    """ Add a column (called TIME_HOURS) based on the data in the TIME column and rounded up 
    the value to the nearest hour.  For example, if the original TIME row said: 
    ‘02/28/2018 05:40:00 PM’ we want ‘2018-02-28 18:00:00’  
    (the change is that 5:40pm was rounded up to 6:00pm and the TIME_HOUR column is 
    actually a proper datetime and not a string).
    """
    from datetime import datetime
    df['TIME_HOURS']=pd.to_datetime(df['TIME']).dt.round('H')
    return df
    raise NotImplementedError()
```

```python
segment_rows = time_to_hours(segment_rows)
```

<!-- grade_id: time_to_hours_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

```python
toto.head()
```

Output:
```
MONTH  DAY_OF_WEEK  HOUR  SEGMENT_ID                    TIME  SPEED  \
0      2            4    17          19  02/28/2018 05:40:00 PM     -1   
1      2            4    17          20  02/28/2018 05:40:00 PM     15   
2      2            4    17          21  02/28/2018 05:40:00 PM     29   
3      2            4    17          22  02/28/2018 05:40:00 PM     28   
4      2            4    17          23  02/28/2018 05:40:00 PM     23   

    STREET DIRECTION FROM_STREET TO_STREET  
0  Pulaski        NB       111th     107th  
1  Pulaski        NB       107th     103rd  
2  Pulaski        NB       103rd      99th  
3  Pulaski        NB        99th      95th  
4  Pulaski        NB        95th      91st
```

<!-- grade_id: filter_orientation, points: 0, grade: false, solution: true -->
```python
def filter_orientation(df, traffic_orientation):
    """ Filter the rows according to the traffic orientation
    return a df that is a subset of the original with the desired orientation
    """
    sbst=df[df['DIRECTION']==traffic_orientation]
    return sbst
    raise NotImplementedError()
```

```python
sb = filter_orientation(segment_rows, 'SB')
nb = filter_orientation(segment_rows, 'NB')
```

The sb table should look like [this](assets/sb.png)

<!-- grade_id: filter_orientation_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell

```

<!-- grade_id: select_column, points: 0, grade: false, solution: true -->
```python
def select_column(df, column_name):
    """ Select a column from the df
    return a series with the desired column
    """
    clm=df[column_name]
    return clm
    raise NotImplementedError()
```

<!-- grade_id: select_column_test, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell

```

```python
# we're going to remove speeds of -1 (no data)
sb = sb[sb.SPEED > -1]
nb = nb[nb.SPEED > -1]
```

```python
sb.head()
```

Output:
```
MONTH  DAY_OF_WEEK  HOUR  SEGMENT_ID                    TIME  SPEED  \
113295     11            5     5          59  11/29/2018 05:50:23 AM     27   
49803       6            6     7          59  06/29/2018 07:50:26 AM     31   
120861     12            2     6          59  12/31/2018 06:50:34 AM     32   
103389     10            6    22          59  10/26/2018 10:50:19 PM     31   
102453     10            6    10          59  10/26/2018 10:50:20 AM     37   

         STREET DIRECTION FROM_STREET TO_STREET          TIME_HOURS  
113295  Pulaski        SB       107th     111th 2018-11-29 06:00:00  
49803   Pulaski        SB       107th     111th 2018-06-29 08:00:00  
120861  Pulaski        SB       107th     111th 2018-12-31 07:00:00  
103389  Pulaski        SB       107th     111th 2018-10-26 23:00:00  
102453  Pulaski        SB       107th     111th 2018-10-26 11:00:00
```

```python
alt.data_transformers.disable_max_rows()
alt.Chart(sb.append(nb)).mark_rect().encode(
    x='month(TIME_HOURS):T',
    y='FROM_STREET:N',
    color='mean(SPEED):Q',
    facet='DIRECTION:N'
).properties(
    width=300,
    height=400
)

```

Output:
```
alt.Chart(...)
```

### 2.5 (Bonus)  Traffic heatmap visualization (up to 2 points)
Looking at the visualization above (the one showing Northbound versus Southbound facets), what domain/abstract tasks are fulfilled by this visualization? List at least one domain task and the corresponding abstract task.

<!-- grade_id: complex_heatmap_vis, points: 0, grade: true, solution: true -->
Domain task: We are comparing between the 2 Pulaski road directions (North and South) in terms of how the mean speed value changes over time for each street. 

Abstract task: Comparing 2 categorical variables. For each one, we are finding the change over time in the correlation between 1 quantitative variable and 1 categorical variable.

### PART B: Crashes (25 points)
For parts 2.6 and 2.7 we will use the Crashes dataset. This dataset contains crash entries recording the time of the accident, the street, and the street number where the accident occurred. You will work with accidents recorded on Pulaski Road


```python
crashes = read_csv('assets/Traffic.Crashes.csv.gz')
crashes_pulaski = crashes[crashes.STREET_NAME == 'PULASKI RD']
```

```python
nono=crashes_pulaski
nono.head()
```

Output:
```
RD_NO CRASH_DATE_EST_I              CRASH_DATE  POSTED_SPEED_LIMIT  \
6    JC100005              NaN  12/31/2018 11:45:00 PM                  25   
51   JB574321              NaN  12/31/2018 09:20:00 PM                  30   
126  JB574112              NaN  12/31/2018 05:30:00 PM                  30   
133  JC111739              NaN  12/31/2018 05:30:00 PM                  30   
240  JB573809              NaN  12/31/2018 01:30:00 PM                  30   

    TRAFFIC_CONTROL_DEVICE      DEVICE_CONDITION WEATHER_CONDITION  \
6           TRAFFIC SIGNAL  FUNCTIONING PROPERLY           UNKNOWN   
51          TRAFFIC SIGNAL  FUNCTIONING PROPERLY              RAIN   
126            NO CONTROLS           NO CONTROLS              RAIN   
133            NO CONTROLS           NO CONTROLS              RAIN   
240            NO CONTROLS           NO CONTROLS              RAIN   

         LIGHTING_CONDITION      FIRST_CRASH_TYPE  \
6    DARKNESS, LIGHTED ROAD               TURNING   
51   DARKNES
```

### 2.6 Calculate summary statistics for grouped streets (15 points)
- Group the streets every 300 units (street numbers). Hint: You can use the `pd.cut` function
- Calculate the number of accidents (count rows) and the total of injuries (sum injuries total) for each of these 300-chunk road segments. Do this *for each direction*.

Complete bin_crashes and calculate_group_aggregates functions for this

<!-- grade_id: bin_crashes, points: 0, grade: false, solution: true -->
```python
def bin_crashes(df):
    """ Assign each crash instance a category (bin) every 300 house number units starting from 0
    Return a new dataframe with a column called BIN where each value is the start of the bin
    i.e. 0 is the label for records with street number n, where 1 <= n <= 300
    300 is the label for records with n at 301 <= n <= 600, and so on.
    """
    st_bins=np.arange(0,df['STREET_NO'].max()+300,300)
    df['BIN']=pd.cut(df['STREET_NO'],bins=st_bins, labels=st_bins[0:-1])
    return df
    raise NotImplementedError()
```

```python
binned_df = bin_crashes(crashes_pulaski)

# sample the values to see what's in your new DF
binned_df.sample(5)[['STREET_NO','BIN']]
```

Output:
```
STREET_NO   BIN
22629        4335  4200
114811       6715  6600
56787        2700  2400
33622        6624  6600
16471        6400  6300
```

A sample of the relevant columns from the table would look something like [this](assets/binned_df.png). We can also create a histogram of street numbers to see which are the most prevalent. It should look something like [this](assets/street_no.png).

```python
# create this vis
alt.Chart(binned_df).mark_bar().encode(
    alt.X('BIN'),
    alt.Y('count()')
)
```

Output:
```
alt.Chart(...)
```

<!-- grade_id: bin_crashes_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

<!-- grade_id: calculate_group_aggregates, points: 0, grade: false, solution: true -->
```python
def calculate_group_aggregates(df):
    """ 
    There are *accidents* and *injuries* (could be 0 people got hurt, could be more). 
    There’s one row per accident at the moment, so we want to know how many accidents 
    happened in each bin (this will be the count) and how many injuries (which will be the sum). 
    
    Return a df with the count of accidents in a column named 'ACCIDENT_COUNT' (how many accidents happened in each 
    bin (the count) and how many injuries (the sum) in a column named 'INJURIES_SUM'
    
    Replace NaN with 0
    """
    new=df.groupby(['BIN','STREET_DIRECTION']).agg({'BIN':len,'INJURIES_TOTAL':np.sum})
    new.replace(np.nan,0,inplace=True)
    new.columns=['ACCIDENT_COUNT','INJURIES_SUM']
    new.reset_index(inplace=True)
    return new
    raise NotImplementedError()
```

```python
aggregates = calculate_group_aggregates(binned_df)

# check the data
#aggregates.head(15)

aggregates.sample(15)
```

Output:
```
BIN STREET_DIRECTION  ACCIDENT_COUNT  INJURIES_SUM
0       0                N            52.0          18.0
2     300                N            37.0          19.0
8    1200                N            76.0          36.0
47   6900                S            75.0           9.0
57   8400                S            52.0          12.0
76  11400                N             0.0           0.0
45   6600                S            91.0          11.0
26   3900                N            80.0          15.0
18   2700                N            69.0           8.0
49   7200                S            37.0           5.0
24   3600                N            35.0           4.0
29   4200                S           116.0          22.0
40   6000                N             9.0           2.0
11   1500                S            31.0          16.0
10   1500                N            42.0          12.0
```

The table should look like [this](assets/2.6_aggregate_1.png)

<!-- grade_id: calculate_group_aggregates_TEST, points: 5, grade: true, solution: false -->
```python
#hidden tests are within this cell

```

Just for fun, here's a plot of injuries in the North and South directions based on bin. This may also help you debug your code. Depending on whether you removed N/A or if you hardcoded things, you may see slight differences. Here's what it might  [look like](assets/direction_injuries.png)

```python
alt.Chart(aggregates).mark_point().encode(
    alt.Color('STREET_DIRECTION'),
    alt.X('BIN'),
    alt.Y('INJURIES_SUM')
)
```

Output:
```
alt.Chart(...)
```

### 2.7 Sort the street ranges (10 points)
- Sort the dataframe so North streets are in descending order and South streets are in ascending order
- You are provided with a 'sort' arrray that contains this desired order. Use a categorical (pd.Categorial) column to order the dataframe according to this array.

<!-- grade_id: categorical_sorting, points: 0, grade: false, solution: true -->
```python
crashed_range = list(range(0, crashes_pulaski.STREET_NO.max()+1000, 300))
sort = ['N ' + str(s) for s in crashed_range[::-1]] + ['S ' + str(s) for s in crashed_range]
def categorical_sorting(df, sort):
    """ Create a column called ORDER_LABEL that contains a concatenation of the street direction and the street range
    Set the sort order of this column to the provided sort array (the elements of this column should be in the same order
    of the array)
    Sort the dataframe by this column
    """
    df['ORDER_LABEL']=pd.Categorical(df[['STREET_DIRECTION','BIN']].astype(str).agg(' '.join,axis=1 ),categories=sort)
    df.sort_values(by='ORDER_LABEL',inplace=True)
    return df
    raise NotImplementedError()
```

```python
sorted_groups = categorical_sorting(aggregates, sort)

# check the values
sorted_groups.sample(15)
```

Output:
```
BIN STREET_DIRECTION  ACCIDENT_COUNT  INJURIES_SUM ORDER_LABEL
3     300                S            49.0          17.0       S 300
66   9900                N             0.0           0.0      N 9900
47   6900                S            75.0           9.0      S 6900
75  11100                S            13.0           4.0     S 11100
36   5400                N             9.0           2.0      N 5400
67   9900                S             6.0           2.0      S 9900
17   2400                S            58.0          17.0      S 2400
62   9300                N             0.0           0.0      N 9300
20   3000                N            87.0          10.0      N 3000
40   6000                N             9.0           2.0      N 6000
53   7800                S            34.0           8.0      S 7800
72  10800                N             0.0           0.0     N 10800
68  10200                N             0.0           0.0     N 10200
49   7200                S            37
```

```python
sorted_groups.sort_values('BIN').tail()
```

Output:
```
BIN STREET_DIRECTION  ACCIDENT_COUNT  INJURIES_SUM ORDER_LABEL
73  10800                S            13.0           1.0     S 10800
75  11100                S            13.0           4.0     S 11100
74  11100                N             0.0           0.0     N 11100
76  11400                N             0.0           0.0     N 11400
77  11400                S             8.0           2.0     S 11400
```

The table should look like [this](assets/sorted_groups.png)

<!-- grade_id: categorical_sorting_test, points: 10, grade: true, solution: false -->
```python
#hidden tests are within this cell

```

Again, just for kicks, let's see where injuries happen. We're going to color bars by the bin and preserve our ascending/descending visualization. We can probably imagine other (better) ways to visualize this data, but this may be useful for you to debug. The visualization should look something like [this](assets/order_injuries.png)

```python
alt.Chart(sorted_groups).mark_bar().encode(
    alt.X('ORDER_LABEL:O', sort=sort),
    alt.Y('INJURIES_SUM:Q'),
    alt.Color('BIN:Q')
).properties(
    width=400
)
```

Output:
```
alt.Chart(...)
```

Ok, let's actually make a useful visualization using some of the dataframes we've created. As a bonus, we're going to ask you what you would use this for.

```python
# to make the kind of chart we are interested in we're going to build it out of three different charts and
# put them together at the end

# this is going to be the left chart
bar_sorted_groups = sorted_groups[['ACCIDENT_COUNT','INJURIES_SUM']].unstack().reset_index() \
    .rename({'level_0':'TYPE','level_1':'SPEED',0:'COUNT'},axis=1)

a = alt.Chart(bar_sorted_groups).mark_bar().transform_filter(alt.datum.TYPE == 'ACCIDENT_COUNT').encode(
    x=alt.X('COUNT:Q',sort='descending'),
    y=alt.Y('SPEED:O',axis=None),
    color=alt.Color('TYPE:N', 
                    legend=None,
                    scale=alt.Scale(domain=['ACCIDENT_COUNT', 'INJURIES_SUM'],
                                    range=['blue', 'orange']))
).properties(
    title='ACCIDENT_COUNT',
    width=300,
    height=600
)

# middle "chart" which actually won't be a chart, just a bunch of labels
b = alt.Chart(bar_sorted_groups).mark_bar().transform_filter(alt.datum.TYPE == 'ACCIDENT_COUNT').encode(
    y=alt.Y('SPEED:O', axis=None),
    text=alt.Text('SPEED:Q')
).mark_text().properties(title='SPEED',
                         width=20,
                         height=600)

# and the right most chart
c = alt.Chart(bar_sorted_groups).mark_bar().transform_filter(alt.datum.TYPE == 'INJURIES_SUM').encode(
    x='COUNT:Q',
    y=alt.Y('SPEED:O',axis=None),
    color=alt.Color('TYPE:N', 
                    legend=None,
                    scale=alt.Scale(domain=['ACCIDENT_COUNT', 'INJURIES_SUM'],
                                    range=['blue', 'orange']))
).properties(
    title='INJURIES_SUM',
    width=300,
    height=600
)

# put them all together 

a | b | c

```

Output:
```
alt.HConcatChart(...)
```

## 2.8 (Bonus) Accident barchart visualization (up to 2 points)
Looking at the visualization we generated above (part 2.7), what domain/abstract tasks are fulfilled by this visualization? List at least one domain task and the corresponding abstract task. 

<!-- grade_id: complex_barchart_vis, points: 0, grade: true, solution: true -->
Domain task: Considering Pulaski Road. For each pair of (road direction, chunk of 300 units), we are measuring both the number of accidents and the total injuries.

Abstract task: The correlation between 1 geographic variable and 2 quantitative variables.